Lab: Implement Entropy & K-L Divergence
====

By The End Of This Lab You Should Be Able To:
----

- Write Python code to calculate the entropy of a discrete distribution
- Write Python code to calculate the K-L Divergence between discrete distributions

dit package
----

dit is a Python package for information theory.

[RTFM](http://docs.dit.io/en/latest/)

```python
reset -fs
```

```python
import sys
import subprocess

try:
    import dit
except ImportError:
    import pip
    subprocess.check_call([sys.executable, '-m', 'pip', 'install', 'dit'])
    import dit
```

```python
# Setup for dit package
outcomes = "🐶 👹 🐯 🐲".split() # Define discrete Random Variable
outcome_probabilities = [0.20, 0.30, 0.25, 0.25] # Created weighted outcomes
assert sum(outcome_probabilities) == 1 # Sanity check
d = dit.Distribution(outcomes, outcome_probabilities) # Create instance

# Let's check it out
print(d)
print()
print(f"The probability of getting a {outcomes[0]} is: {d[outcomes[0]]}")
print(f"The probability of getting a {outcomes[0]} and {outcomes[1]} is: {d.event_probability([outcomes[0], outcomes[1]])}")
```

Output:
```
Class:          Distribution
Alphabet:       ('🐯', '🐲', '🐶', '👹') for all rvs
Base:           linear
Outcome Class:  str
Outcome Length: 1
RV Names:       None

x   p(x)
🐯   0.25
🐲   0.25
🐶   0.2
👹   0.3

The probability of getting a 🐶 is: 0.2
The probability of getting a 🐶 and 👹 is: 0.5
```

```python
print(f"The Shannon entropy of this pmf is: {dit.shannon.entropy(d):.4f}")
```

Output:
```
The Shannon entropy of this pmf is: 1.9855
```

<!-- grade_id: cell-d80194fc50e3b03b, points: 0, grade: false, solution: true -->
```python
from typing import List

def shannon_entropy(outcome_probabilities: List[float]) -> float:
    """Implement Shannon entropy function.
    You may use any math or numpy method. 
    You may NOT use any other package, including `scipy.stats.entropy`
    If you use another package, you'll get zero points.
    """ 
    import math
    import numpy as np

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return H
```

<!-- grade_id: cell-7ff458e86e415763, points: 2, grade: true, solution: false -->
```python
"""
2 points
Test code for the 'shannon_entropy' function. 
This cell should NOT give any errors when it is run.
"""
from math import isclose

outcomes = '🧟 🧙'.split()
outcome_probabilities = [0, 1]
assert sum(outcome_probabilities) == 1 # Sanity check
d = dit.Distribution(outcomes, outcome_probabilities)
assert isclose(dit.shannon.entropy(d), shannon_entropy(outcome_probabilities))

outcomes = '🧖 🧗'.split()
outcome_probabilities = [0.2, 0.8]
assert sum(outcome_probabilities) == 1 # Sanity check
d = dit.Distribution(outcomes, outcome_probabilities)
assert isclose(dit.shannon.entropy(d), shannon_entropy(outcome_probabilities))

outcomes = "🐶 👹 🐯 🐲".split()
outcome_probabilities = [0.35, 0.15, 0.25, 0.25]
assert sum(outcome_probabilities) == 1 # Sanity check
d = dit.Distribution(outcomes, outcome_probabilities)
assert isclose(dit.shannon.entropy(d), shannon_entropy(outcome_probabilities))
```

<!-- grade_id: cell-4cf71a55c9e1c9ab, points: 0, grade: false, solution: true -->
```python
def decrease_entropy(probabilities_orginal: List[float]) -> List[float]:
    "Manually change the individual entries to decrease the overall entropy."
    # YOUR CODE HERE
    raise NotImplementedError()
    return probabilities_updated
```

<!-- grade_id: cell-c7ae99e89efedd66, points: 1, grade: true, solution: false -->
```python
"""
1 point
Test code for the 'decrease_entropy' function. 
This cell should NOT give any errors when it is run.
"""

outcomes = "🐶 👹 🐯 🐲".split()
probabilities_orginal = [0.35, 0.15, 0.25, 0.25]
probabilities_updated = decrease_entropy(probabilities_orginal)
assert sum(probabilities_updated) == 1 # Sanity check
d = dit.Distribution(outcomes, probabilities_updated)
assert dit.shannon.entropy(d) < 1.94
```

<!-- grade_id: cell-2e243b1ac3126a3c, points: 0, grade: false, solution: true -->
```python
def increase_entropy(probabilities_orginal: List[float]) -> List[float]:
    "Manually change the individual entries to increase the overall entropy."
    # YOUR CODE HERE
    raise NotImplementedError()
    return probabilities_updated
```

<!-- grade_id: cell-94dafe9da4b23ef2, points: 1, grade: true, solution: false -->
```python
"""
1 point
Test code for the 'increase_entropy' function. 
This cell should NOT give any errors when it is run.
"""

outcomes = "🐶 👹 🐯 🐲".split()
probabilities_orginal = [0.35, 0.15, 0.25, 0.25]
probabilities_updated = increase_entropy(probabilities_orginal)
assert sum(probabilities_updated) == 1 # Sanity check
d = dit.Distribution(outcomes, probabilities_updated)
assert dit.shannon.entropy(d) > 1.94
```

<!-- grade_id: cell-2bda54533b8398eb, points: 0, grade: false, solution: true -->
```python
def maximum_entropy(probabilities_orginal: List[float]) -> List[float]:
    "Manually change the individual entries to maximize the overall entropy."
    # YOUR CODE HERE
    raise NotImplementedError()
    return probabilities_updated
```

<!-- grade_id: cell-4493bf867196072b, points: 1, grade: true, solution: false -->
```python
"""
1 point
Test code for the 'maximum_entropy' function. 
This cell should NOT give any errors when it is run.
"""

outcomes = "🐶 👹 🐯 🐲".split()
probabilities_orginal = [0.35, 0.15, 0.25, 0.25]
probabilities_updated = maximum_entropy(probabilities_orginal)
assert sum(probabilities_updated) == 1 # Sanity check
d = dit.Distribution(outcomes, probabilities_updated)
assert dit.shannon.entropy(d) == 2
```

<!-- grade_id: cell-514f4738a324a949, points: 0, grade: false, solution: true -->
```python
def minimum_entropy(probabilities_orginal: List[float]) -> List[float]:
    "Manually change the individual entries to minimumize the overall entropy."
    # YOUR CODE HERE
    raise NotImplementedError()
    return probabilities_updated
```

<!-- grade_id: cell-c66630817b2dc21d, points: 1, grade: true, solution: false -->
```python
"""
1 point
Test code for the 'minimum_entropy' function. 
This cell should NOT give any errors when it is run.
"""

outcomes = "🐶 👹 🐯 🐲".split()
probabilities_orginal = [0.35, 0.15, 0.25, 0.25]
probabilities_updated = minimum_entropy(probabilities_orginal)
assert sum(probabilities_updated) == 1 # Sanity check
d = dit.Distribution(outcomes, probabilities_updated)
assert dit.shannon.entropy(d) == 0
```

-----
K-L Divergence 
------

We have a Random Variable that models two lunch states:

1. Candy, aka 🍬
1. Sushi, aka 🍣

X ={🍭, 🍣}

There are two childern:

1. Patel, aka p, 
1. Quincy, aka q

Each child has preference for one state for lunch, p has prefernce r and q has and preference s. We can model those childern as two different Bernoulli distributions over those states:

p(🍭) = 1-r  
p(🍣) = r

q(🍭) = 1-s   
q(🍣) = s 

<!-- grade_id: cell-8aa40b36fdf2c217, points: 0, grade: false, solution: true -->
```python
def kl_divergence_two_states(r, s)-> float:
    """Implement K-L Divergence for two states.
    You may use any math or numpy method. 
    You many NOT use any other package, including `scipy.special.kl_div`
    If you use another package, you'll get zero points.
    """
    import math
    import numpy as np

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return kl_div
```

<!-- grade_id: cell-1dc5a0be61f4b059, points: 5, grade: true, solution: false -->
```python
"""
5 points
Test code for the 'kl_divergence_two_states' function. 
This cell should NOT give any errors when it is run.
"""

r = 1/2 
s = 1/4
kl_div = kl_divergence_two_states(r, s)
assert f"{kl_div:.6f}" == '0.207519'

r = 1/4 # NOTE: The parameters are swapped
s = 1/2 # NOTE: The parameters are swapped
kl_div = kl_divergence_two_states(r, s)
assert f"{kl_div:.6f}" != '0.207519' # NOTE: Not symmetrical
```

Now Patel and Quincy have to decide about dinner which has more options.

There is:

1. Noodles, aka 🍜
2. Steak, aka 🥩
3. Crab, aka 🦀

```python
# Store states and probabilities in a dict
p = {'🍜':.5,     '🥩':.25,   '🦀':.25}
q = {'🍜':.58333, '🥩':.1666, '🦀':.25}

# Kids ❤️ noodles
```

```python
import pandas as pd

%matplotlib inline
```

```python
pd.Series(p).plot(kind='bar', title="Patel's preferences");
```

```python
pd.Series(q).plot(kind='bar', title="Quincey's preferences");
```

<!-- grade_id: cell-ea2c0c82f2cc1b3c, points: 0, grade: false, solution: true -->
```python
def kl_divergence_two_discrete_distrubtions(p, q)-> float:
    """Implement K-L Divergence for two discrete distributions.
    You may use any math or numpy method. 
    You many NOT use any other package, including `scipy.special.kl_div`
    If you use another package, you'll get zero points.
    """
    import math
    import numpy as np
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return kl_div
```

<!-- grade_id: cell-119b629d46e0cebc, points: 5, grade: true, solution: false -->
```python
"""
5 points
Test code for the 'kl_divergence_two_discrete_distrubtions' function. 
This cell should NOT give any errors when it is run.
"""

kl_div = kl_divergence_two_discrete_distrubtions(p, q)
assert f"{kl_div:.6f}" == '0.035193'

kl_div = kl_divergence_two_discrete_distrubtions(q, p) # NOTE: The parameters are swapped
assert f"{kl_div:.6f}" != '0.035193' # NOTE: Not symmetrical
```

<br>
<br> 
<br>

----