Lab: Random Forestâ„¢
====

By The End Of This Lab You Should Be Able To:
----

- Fit Random Forestâ„¢ to variety of datasets
- Download data from the web
- Munge categorical data into a format amendable to Machine Learning


Parts
------


1. Toy dataset - Moons ðŸŒ›
1. Built-in dataset - Wine ðŸ·
1. "Wild" dataset - Mushroom ðŸ„

Part 1a: Fit Random Forestâ„¢ to Moon dataset
------

```python
reset -fs
```

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
import seaborn as sns
import sklearn

import warnings
warnings.filterwarnings('ignore')

palette = "Dark2"
%matplotlib inline
```

```python
from sklearn.datasets import make_moons
```

```python
noisy_moons = make_moons(n_samples=10_000, noise=.25, random_state=42)

sns.scatterplot(x=noisy_moons[0][:, 0],
                y=noisy_moons[0][:, 1],
                hue=noisy_moons[1],
                palette=palette);
```

```python
from sklearn.model_selection import train_test_split

# Load and split the data
X_train, X_test, y_train, y_test = train_test_split(noisy_moons[0], noisy_moons[1], test_size=0.3, random_state=42)
```

```python
from sklearn.ensemble import RandomForestClassifier
```

<!-- grade_id: cell-664ed1ccd2613e44, points: 0, grade: false, solution: true -->
```python
"""
Fit a RandomForestClassifier to this data. Find accuracy on test data.

NOTE: Just write code in this cell. Do not create a function.
"""

# YOUR CODE HERE
raise NotImplementedError()

print(f"{rf_moon_acc:.4f}")
```

<!-- grade_id: cell-48aeec842a9d1000, points: 2, grade: true, solution: false -->
```python
"""
2 points
Test code for the cell above.
This cell should NOT give any errors when it is run.
"""
from math import isclose

assert rf_moon.__class__.__name__.split('.')[-1] == 'RandomForestClassifier'
assert isclose(rf_moon_acc, 0.9336666666666666)
```

Part 1b: Grid Search Random Forestâ„¢ to Moon dataset
------

Hint: Is there a Random Forestâ„¢ hyperparameter that could spend up training?

<!-- grade_id: cell-5f0e8fa34dea08af, points: 0, grade: false, solution: true -->
```python
"""
Write code to conduct Grid Search on RandomForestClassifier to find best number of trees and best criterion.

NOTE: Just write code in this cell. Do not create a function.
"""

# YOUR CODE HERE
raise NotImplementedError()

print(f"{gs_rf_acc:.4f}")
```

<!-- grade_id: cell-9b73c1a1f260d93a, points: 4, grade: true, solution: false -->
```python
"""
4 points
Test code for the cell above.
There are hidden tests. They will make sure best_params_ have the optimal values.
This cell should NOT give any errors when it is run.
"""
from math import isclose

best_params = gs.best_params_
assert set(best_params.keys()) == {'criterion', 'n_estimators'}
assert isclose(gs_rf_acc, 0.9353333333333333)

```

Part 2: Fit Random Forestâ„¢ to Wine dataset
------

```python
from sklearn.datasets import load_wine
from sklearn.model_selection import train_test_split

data = load_wine()
X_train, X_test, y_train, y_test = train_test_split(data.data, data.target, test_size=0.3, random_state=42)
```

<!-- grade_id: cell-5d4f709b1d3c8e9a, points: 0, grade: false, solution: true -->
```python
"""
Fit a RandomForestClassifier to this data. Find accuracy on test data.

NOTE: Just write code in this cell. Do not create a function.
"""

# YOUR CODE HERE
raise NotImplementedError()
print(f"{rf_wine_acc:.4f}")
```

<!-- grade_id: cell-d5d99add6e5604de, points: 2, grade: true, solution: false -->
```python
"""
2 points
Test code for the cell above.
This cell should NOT give any errors when it is run.
"""
from math import isclose

assert rf_wine.__class__.__name__.split('.')[-1] == 'RandomForestClassifier'
assert isclose(rf_wine_acc, 0.9259259259259259)
```

<!-- grade_id: cell-79085de7f6004f82, points: 0, grade: false, solution: true -->
```python
"""
Create a sorted list of features by importance.

NOTE: Just write code in this cell. Do not create a function.
"""

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ae9151666236836a, points: 5, grade: true, solution: false -->
```python
"""
5 points
Test code for the cell above.
There are hidden tests. They will make sure features_most_important_to_least are correctly ordered
This cell should NOT give any errors when it is run.
"""

features_most_important_to_least[0] == 'flavanoids'
features_most_important_to_least[-1] == 'nonflavanoid_phenols'

```

Part 3: Fit Random Forestâ„¢ to Mushroom dataset
------

Descriptions of hypothetical samples mushrooms. Each species is identified as edible (e) or poisonous (p).

Learn more:

- https://archive.ics.uci.edu/ml/datasets/mushroom
- https://www.kaggle.com/uciml/mushroom-classification/version/1

```python
import os

from urllib.request import urlretrieve

path_remote = "https://archive.ics.uci.edu/ml/machine-learning-databases/mushroom/"
path_local = "."
filename = "agaricus-lepiota.data"
```

<!-- grade_id: cell-a3f25e54e48fc92b, points: 0, grade: false, solution: true -->
```python
def download_data(path_remote: str, path_local: str, filename: str) -> None:
    """Write code that retrieves the data from url.
    Make sure to check if file is local, if not then retrieve it.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9e7b459e9195818c, points: 2, grade: true, solution: false -->
```python
"""
2 points
Test code for the 'download_data' function. 
This cell should NOT give any errors when it is run.
"""

local_file = os.path.join(path_local, filename)
if os.path.exists(local_file):
    os.remove(local_file)
    
download_data(path_remote, path_local, filename)

assert os.path.exists(local_file)
```

<!-- grade_id: cell-4327521a88fb5854, points: 0, grade: false, solution: true -->
```python
"""
Write code to do the following:
1. Load data with Pandas
2. Separate targets / labels
3. Convert targets / labels to dummy encoding
4. Convert every categorical feature with one hot encoding

Hint: Remove unneeded columns from your dataframe as you go.

NOTE: Just write code in this cell. Do not create a function.
"""

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-11793b6c57d3b00c, points: 10, grade: true, solution: false -->
```python
"""
10 points
Test code for the cell above.
This cell should NOT give any errors when it is run.
"""

assert type(labels) == pd.core.series.Series
assert labels.shape == (8123,)
assert labels[0] == 0
assert labels[42] == 1
assert labels[8122] == 0

assert type(features) == pd.core.frame.DataFrame
assert features.shape == (8123, 117)
assert features.x_b.iloc[0] == 0
assert features.u_l.iloc[8122] == 1
```

```python
X_train, X_test, y_train, y_test = train_test_split(features, labels, test_size=0.3, random_state=42)
```

<!-- grade_id: cell-866f7c557da11b80, points: 0, grade: false, solution: true -->
```python
"""
Fit a RandomForestClassifier to this data. Find accuracy on test data.

NOTE: Just write code in this cell. Do not create a function.
"""
# YOUR CODE HERE
raise NotImplementedError()

print(f"{rf_mushroom_acc:.4f}")
```

<!-- grade_id: cell-7a5d0bf123c0b1a9, points: 2, grade: true, solution: false -->
```python
"""
2 points
Test code for the cell above.
This cell should NOT give any errors when it is run.
"""

assert rf_mushroom.__class__.__name__.split('.')[-1] == 'RandomForestClassifier'
assert rf_mushroom_acc == 1.0000
```

<!-- grade_id: cell-29d87ba05cf3c7c8, points: 0, grade: false, solution: true -->
```python
"""
Fit a Naive Bayes to this data. Find accuracy on test data.

Hint: You must import the correct implementation of Naive Bayes from scikit-learn. 

NOTE: Just write code in this cell. Do not create a function.
"""
# YOUR CODE HERE
raise NotImplementedError()

print(f"{nb_mushroom_acc:.4f}")
```

<!-- grade_id: cell-70cc296a04f113e9, points: 2, grade: true, solution: false -->
```python
"""
2 points
Test code for the cell above.
There are hidden tests. They will make sure you used the correct version of Naive Bayes.
This cell should NOT give any errors when it is run.
"""
from math import isclose

assert isclose(nb_mushroom_acc, 0.9376282314320886)
assert set(nb_mushroom.classes_) == {0, 1} 
 
```

Bonus Cartoon
------

<center><img src="https://imgs.xkcd.com/comics/mushrooms.png" width="75%"/></center>


<br>
<br> 
<br>

----