Lab: k-nearest neighbors (k-NN) with scikit-learn
===

By The End Of This Session You Should Be Able To:
----

- Visually interpret data
- Fit KNN model with scikit-learn
- Explain how features effects modeling
- Explain how changing the number of neighbors effects modeling

```python
reset -fs
```

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
import seaborn as sns
import sklearn

import warnings
warnings.filterwarnings('ignore')

palette = "Dark2"
%matplotlib inline
```

Iris Data
-----

<center><img src="http://s5047.pcdn.co/wp-content/uploads/2015/04/iris_petal_sepal.png" width="75%"/></center>

```python
# Load data
iris_sns = sns.load_dataset("iris")
iris_sns.head(n=3)
```

```python
# Always a good habit to manually inspect the tail
iris_sns.tail(n=3)
```

```python
# Pretty plot
sns.pairplot(iris_sns, hue='species', size=2.5, palette=palette);
```

Which two dimensions would provide for the __best__ seperation between the three classes?

<details><summary>
Click here for the solution…
</summary>
petal_length  and petal_width 
</details>

Which two dimensions would provide for the __worst__ seperation between the three classes?

<details><summary>
Click here for the solution…
</summary>
sepal_length and sepal_width
</details>

-----
Modeling
-----

We are going to take a model comparison approach in the class.

First, we are going to fit baseline model. Then try to improve the models from there. We might also make things worse.

```python
# Load data for modeling
from sklearn.datasets import load_iris

iris = load_iris()
```

```python
# scikit-learn datasets are dicts
iris.keys()
```

```python
print(iris.DESCR)
```

<!-- grade_id: cell-d80194fc50e3b03b, points: 0, grade: false, solution: true -->
```python
def select_two_worst_features(iris: sklearn.utils.Bunch) -> np.array:
    "Select the two WORST features/columns from iris"
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7ff458e86e415763, points: 1, grade: true, solution: false -->
```python
"""
1 point
Test code for the 'select_two_worst_features' function. 
This cell should NOT give any errors when it is run.
"""

assert select_two_worst_features(iris).shape == (150, 2)
assert list(select_two_worst_features(iris)[0]) == [5.1, 3.5]
assert list(select_two_worst_features(iris)[-1]) == [5.9, 3. ]
```

Let's check out the documention for [KNeighborsClassifier](http://scikit-learn.org/stable/modules/generated/sklearn.neighbors.KNeighborsClassifier.html)

```python
from sklearn.neighbors import KNeighborsClassifier
```

```python
# Create an instance of the model from the class
knn = KNeighborsClassifier(n_neighbors=3)
```

```python
# Define the input data and labels
X = select_two_worst_features(iris)
y = iris.target

# Fit the model
knn.fit(X, y)
```

```python
# Predict a within sample datapoint
i = 100 # Index could also be a random integer 
y_predicted = knn.predict(X[i].reshape(1, -1))
y_predicted
```

```python
y_actual = y[i]
y_actual
```

```python
# Pretty plot the selected data
scatter = sns.scatterplot(x=X[:, 0],
                          y=X[:, 1],
                          hue=y,     #'hue' will color code each group
                          palette=palette,
                         );


# Plot data point
scatter.plot(X[i, 0], 
            X[i, 1], 
            color='red', 
            marker='*');
```

You have just fit your a model with sci-kit learn  💥

```python
# Predict an out of sample datapoint
noise = -.35 # Noise could also be a random float
new_data = X[i].reshape(1, -1)+noise
y_predicted = knn.predict(new_data)
y_predicted
```

```python
# Pretty plot the selected data
scatter = sns.scatterplot(x=X[:, 0],
                          y=X[:, 1],
                          hue=y,     #'hue' will color code each group
                          palette=palette,
                         );


# Plot data point
scatter.plot(new_data[0][0], 
            new_data[0][1], 
            color='red', 
            marker='*');
```

```python
# Predict the response for every within sample data
# NOTE: We are not doing train / test split. Splitting the data is pointless with k NN because it just memorized the observed data.
y_predicted = knn.predict(X)
```

```python
from sklearn.metrics import accuracy_score
```

```python
accuracy = accuracy_score(y_predicted, y)
print(f'Model accuracy {accuracy:.2%}')
```

Why do you think is the accuracy so low? (ungraded)

<!-- grade_id: cell-a41ec0c70c3d9171, points: 0, grade: false, solution: true -->
```python
def fit_knn(X: np.array, y: np.array, n_neighbors: int=3) -> float:
    "Fit a KNN model, returning accuracy. Use code above an example."
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2c15b5ce3e489b5a, points: 1, grade: true, solution: false -->
```python
"""
1 point
Test code for the 'fit_KNN' function with 2 worst features
This cell should NOT give any errors when it is run.
"""

assert round(fit_knn(X, y, n_neighbors=3), 4) == .8533
```

<!-- grade_id: cell-d744dae2cf569db0, points: 0, grade: false, solution: true -->
```python
def select_two_best_features(iris: sklearn.utils.Bunch) -> np.array:
    "Select the two BEST features/columns from iris and return as numpy.ndarray"
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c0d084fefedea43b, points: 1, grade: true, solution: false -->
```python
"""
1 point
Test code for the 'select_two_worst_features' function. 
This cell should NOT give any errors when it is run.
"""

assert select_two_best_features(iris).shape == (150, 2)
assert list(select_two_best_features(iris)[0]) == [1.4, 0.2]
assert list(select_two_best_features(iris)[-1]) == [5.1, 1.8]
```

```python
X = select_two_best_features(iris)
```

<!-- grade_id: cell-f57f9b9f7f82fbce, points: 1, grade: true, solution: false -->
```python
"""
1 point
Test code for the 'fit_KNN' function with 2 best features.
This cell should NOT give any errors when it is run.
"""

assert round(fit_knn(X, y, n_neighbors=3), 4) == 0.98
```

```python
X = iris.data  # Use all features
```

<!-- grade_id: cell-8da2ab30adba2afc, points: 1, grade: true, solution: false -->
```python
"""
1 point
Test code for the 'fit_KNN' function with all features.
This cell should NOT give any errors when it is run.
"""

assert round(fit_knn(X, y, n_neighbors=3), 4) == 0.96
```

Why do you think the accuracy go down when we have more features? (ungraded)

-----
Explore how the number of neighbors impacts modeling
-----

```python
X = iris.data  # Use all features

print(f'{"# of Neighbors"} | {"Model accuracy":>8}')
print('-'*35)
for n_neighbors in range(1, 16):
    accuracy = fit_knn(X, y, n_neighbors)
    print(f'{n_neighbors:^14} | {accuracy:>10.2%}')
```

Why is accuracy 100% with 1 neighbor? (ungraded)

Why does accuracy go down then back up as number of neighbors increase? (ungraded)

<br>
<br> 
<br>

----